# Lesson name

Some theory and examples.

```python
print('Hey!')
```

Output:
```
Hey!
```

The cells below show how each task should be designed.

#### TODO: Triangle perimeter

Write a function that determines the perimeter of a triangle using the length of three sides.

<!-- grade_id: cell-b791b685d2eb9689, points: 0, grade: false, solution: true -->
```python
def get_perimeter(a: float, b: float, c: float) -> float:
    """Get triangle perimeter.

    :param a: first triangle side.
    :param b: second triangle side.
    :param c: third triangle side.
    :return: perimeter.
    """
    # Write your code instead of None
    perimeter = None
    ### BEGIN SOLUTION
    perimeter = a + b + c
    ### END SOLUTION
    return perimeter
```

<!-- grade_id: cell-059ec36d3798e24a, points: 100, grade: true, solution: false -->
```python
# Check
perimeter = get_perimeter(3, 4, 5)
print(f'Perimeter of triangle with sides 3, 4, 5: {perimeter}')
### BEGIN HIDDEN TESTS
assert get_perimeter(5, 5, 5) == 15
assert get_perimeter(2, 3, 4) == 9
assert get_perimeter(15, 24, 18) == 57
### END HIDDEN TESTS
```

Output:
```
Perimeter of triangle with sides 3, 4, 5: 12
```

Expected result:

Perimeter of triangle with sides 3, 4, 5: 12

---